```python
pip install nose
```

Output:
```
Requirement already satisfied: nose in /opt/conda/lib/python3.10/site-packages (1.3.7)
Note: you may need to restart the kernel to use updated packages.
```

```python
#Perform logistic regression on diabetes dataset
```

<!-- grade_id: cell-65901a02d7b3d454, points: 0, grade: false, solution: true -->
```python
#Load diabetes dataset here and return shape of dataset
from cz_utility import dataset
# diabetes dataset
df = dataset.load('63ea33b29ac5f900139b8d07')
def returnShape(df):
    # your code here
    return df.shape
    raise NotImplementedError
```

<!-- grade_id: cell-9d4cd5b17d896c66, points: 2, grade: true, solution: false -->
```python
from nose.tools import assert_equal
assert_equal(returnShape(df), (768, 9))
```

```python
df
```

Output:
```
Pregnancies  Glucose  BloodPressure  SkinThickness  Insulin   BMI  \
0              6      148             72             35        0  33.6   
1              1       85             66             29        0  26.6   
2              8      183             64              0        0  23.3   
3              1       89             66             23       94  28.1   
4              0      137             40             35      168  43.1   
..           ...      ...            ...            ...      ...   ...   
763           10      101             76             48      180  32.9   
764            2      122             70             27        0  36.8   
765            5      121             72             23      112  26.2   
766            1      126             60              0        0  30.1   
767            1       93             70             31        0  30.4   

     DiabetesPedigreeFunction  Age  Outcome  
0                       0.627   50        1  
1                       
```

```python
import pandas as pd
import numpy as np

# Calculate correlation coefficients between features and target variable
correlation = df.corr()['Outcome'].abs().sort_values()

# Display correlation coefficients
print(correlation)

# Get the feature(s) with the lowest correlation coefficient(s)
least_correlated_features = correlation[correlation == correlation.min()].index

# Display the least correlated feature(s)
print("Least correlated feature(s):", least_correlated_features)

```

Output:
```
BloodPressure               0.065068
SkinThickness               0.074752
Insulin                     0.130548
DiabetesPedigreeFunction    0.173844
Pregnancies                 0.221898
Age                         0.238356
BMI                         0.292695
Glucose                     0.466581
Outcome                     1.000000
Name: Outcome, dtype: float64
Least correlated feature(s): Index(['BloodPressure'], dtype='object')
```

<!-- grade_id: cell-97b7b8d9c24aa579, points: 0, grade: false, solution: true -->
```python
#separate input and output from dataset and remove unwanted data columns
#take x for input and y for output
def sepAndRemove(df):
    # your code here
    x = df.drop(columns=['Outcome','BloodPressure'])
    y = df['Outcome']
    return x, y
    raise NotImplementedError
```

<!-- grade_id: cell-f7e39f11fcd55f39, points: 2, grade: true, solution: false -->
```python
trainAndTestData = sepAndRemove(df)
assert_equal(trainAndTestData[0].shape, (768, 7))
assert_equal(trainAndTestData[1].shape, (768,))
```

```python
trainAndTestData[0].shape
```

Output:
```
(768, 7)
```

```python
trainAndTestData[1]
```

Output:
```
0      1
1      0
2      1
3      0
4      1
      ..
763    0
764    0
765    0
766    1
767    0
Name: Outcome, Length: 768, dtype: int64
```

<!-- grade_id: cell-05c4eb3e485c557a, points: 0, grade: false, solution: true -->
```python
#separate training and testing dataset with test_size = 0.3
from sklearn.model_selection import train_test_split
def trainTestSplit(trainAndTestData):
    # your code here
    train_x, test_x, train_y, test_y = train_test_split(trainAndTestData[0], trainAndTestData[1], test_size=0.3,random_state=42)
    return train_x, test_x, train_y, test_y
    raise NotImplementedError
```

```python

```

<!-- grade_id: cell-a7a6e742dbef0d20, points: 2, grade: true, solution: false -->
```python
trainAndTestData=trainTestSplit(sepAndRemove(df))
assert_equal(trainAndTestData[0].shape,(537, 7))
assert_equal(trainAndTestData[1].shape,(231, 7))
assert_equal(trainAndTestData[2].shape,(537,))
assert_equal(trainAndTestData[3].shape,(231,))
```

```python
trainAndTestData[0].shape
```

Output:
```
(537, 7)
```

```python
trainAndTestData[1].shape
```

Output:
```
(231, 7)
```

```python
trainAndTestData[2].shape
```

Output:
```
(537,)
```

```python
trainAndTestData[3].shape
```

Output:
```
(231,)
```

<!-- grade_id: cell-6e973aeb54f8a469, points: 0, grade: false, solution: true -->
```python
#use MinMaxScaler to scale the data to a fixed range
from sklearn.preprocessing import MinMaxScaler
def scalar(trainAndTestData):
    # your code here
    scl = MinMaxScaler()
    trainX_scaled = scl.fit_transform(trainAndTestData[0])
    testX_scaled = scl.transform(trainAndTestData[1])
    return trainX_scaled, testX_scaled, trainAndTestData[2], trainAndTestData[3]
    raise NotImplementedError
```

<!-- grade_id: cell-0bb62bf101293688, points: 2, grade: true, solution: false -->
```python
trainAndTestData=scalar(trainTestSplit(sepAndRemove(df)))
assert_equal(trainAndTestData[0].shape, (537, 7))
assert_equal(trainAndTestData[1].shape, (231, 7))

```

```python
# trainAndTestData[0].shape

trainAndTestData[0]
```

Output:
```
array([[0.05882353, 0.47738693, 0.28571429, ..., 0.3561848 , 0.07771136,
        0.01666667],
       [0.29411765, 0.52763819, 0.46031746, ..., 0.54992548, 0.03458582,
        0.11666667],
       [0.        , 0.67839196, 0.66666667, ..., 0.63040238, 0.12254483,
        0.05      ],
       ...,
       [0.58823529, 0.50753769, 0.58730159, ..., 0.67958271, 0.45175064,
        0.28333333],
       [0.        , 0.70854271, 0.        , ..., 0.6318927 , 0.05422716,
        0.13333333],
       [0.        , 0.6281407 , 0.        , ..., 0.33532042, 0.07856533,
        0.        ]])
```

```python
# trainAndTestData[1].shape


trainAndTestData[1]
```

Output:
```
array([[0.35294118, 0.49246231, 0.52380952, ..., 0.50670641, 0.15029889,
        0.36666667],
       [0.11764706, 0.56281407, 0.50793651, ..., 0.53204173, 0.02988898,
        0.        ],
       [0.11764706, 0.54271357, 0.        , ..., 0.45901639, 0.03415884,
        0.        ],
       ...,
       [0.        , 0.63819095, 0.58730159, ..., 0.54098361, 0.30999146,
        0.03333333],
       [0.35294118, 0.52763819, 0.50793651, ..., 0.45901639, 0.01878736,
        0.26666667],
       [0.29411765, 0.38693467, 0.65079365, ..., 0.53353204, 0.03330487,
        0.23333333]])
```

```python
trainAndTestData[2].shape
```

Output:
```
(537,)
```

```python
trainAndTestData[3].shape
```

Output:
```
(231,)
```

<!-- grade_id: cell-a9178e0adf01003e, points: 0, grade: false, solution: true -->
```python
#perform logistic regression here and return accuracy between predected output for testx and testy
from sklearn.linear_model import LogisticRegression
from sklearn.metrics import accuracy_score
def logisticReg(trainAndTestData, trainy, testy):
    # your code here
    model = LogisticRegression()
    model.fit(trainAndTestData[0], trainy)
    pred_y = model.predict(trainAndTestData[1])
    accuracy = accuracy_score(testy, pred_y)
    print(accuracy)
    return accuracy
    raise NotImplementedError
```

<!-- grade_id: cell-b26830e8ae1d3a06, points: 2, grade: true, solution: false -->
```python
#there might be difference between student answer and faculty answer
trainAndTestData1 = trainTestSplit(sepAndRemove(df))
accuracyScore = logisticReg(scalar(trainTestSplit(sepAndRemove(df))), trainAndTestData1[2], trainAndTestData1[3])
assert_equal(accuracyScore>0.60, True)
```

Output:
```
0.7489177489177489
```

```python

```